```python
# Initialize Otter
import otter
grader = otter.Notebook("hw07.ipynb")
```

# Homework 7: Bike Sharing
 
Exploratory Data Analysis (EDA) and Visualization

(This homework is based on [Data 100 Fall 2022 hw 04](https://github.com/DS-100/fa22/tree/main/hw/hw04).)

<!-- BEGIN QUESTION -->

**Question:** What is your name?

_Type your answer here, replacing this text._

<!-- END QUESTION -->

## Introduction

Bike sharing systems are a new generation of traditional bike rentals where the process of signing up, renting and returning is automated. Through these systems, users are able to easily rent a bike from one location and return them to another. We will be analyzing bike sharing data from Washington D.C. 

In this assignment, you will perform tasks to visualize and explore the bike sharing data. You will also investigate open-ended questions. These open-ended questions ask you to think critically about how the plots you have created provide insight into the data.

For plots, your plots should be *similar* to the given examples. It is okay to have small variations such as color differences or slight variations in scale. 
Your plots should have appropriate titles, axis labels, legends, etc. 
Ask yourself:
> If I directly downloaded the plot and viewed it, would I be able to tell what was being visualized without knowing the question? 

```python
# Don't change this cell; just run it. 
import numpy as np
import seaborn as sns
from datascience import *

# These lines do some fancy plotting magic.
import matplotlib
%matplotlib inline
import matplotlib.pyplot as plt
#plt.style.use('fivethirtyeight')
plt.rcParams['figure.figsize'] = (16,8)
plt.rcParams['figure.dpi'] = 150
sns.set()
```

## 1. Loading and Examining the Bike Sharing Data
The data we are exploring is collected from a bike sharing system in Washington D.C.

The variables are defined as:

Variable       | Description
-------------- | ------------------------------------------------------------------
instant | record index
dteday | date
season | 1. spring <br> 2. summer <br> 3. fall <br> 4. winter
yr | year (0: 2011, 1:2012)
mnth | month ( 1 to 12)
hr | hour (0 to 23)
holiday | whether day is holiday or not
weekday | day of the week
workingday | if day is neither weekend nor holiday
weathersit | 1. clear or partly cloudy <br> 2. mist and clouds <br> 3. light snow or rain <br> 4. heavy rain or snow
temp | normalized temperature in Celsius (divided by 41)
atemp | normalized "feels-like" temperature in Celsius (divided by 50)
hum | normalized percent humidity (divided by 100)
windspeed| normalized wind speed (divided by 67)
casual | count of casual users
registered | count of registered users
cnt | count of total rental bikes including casual and registered  

Load the data below.

```python
bike = Table().read_table("bikeshare.csv")
bike
```

<!-- BEGIN QUESTION -->

**Question 1.1:** What is the granularity of the data (i.e. what does each row represent)?

_Type your answer here, replacing this text._

<!-- END QUESTION -->



```python
# Use this cell for scratch work. If you need to add more cells for scratch work, add them BELOW this cell.
```

<!-- BEGIN QUESTION -->

**Question 1.2:**
For this assignment, we'll be using this data to study bike usage in Washington D.C. Based on the granularity and the variables present in the data, what might some limitations of using this data be? What are two additional data categories/variables that you could collect to address some of these limitations?

_Type your answer here, replacing this text._

<!-- END QUESTION -->



```python
# Use this cell for scratch work. If you need to add more cells for scratch work, add them BELOW this cell.
```

The variables table above tells us

Variable       | Description
-------------- | ------------------------------------------------------------------
holiday | whether day is holiday or not
weekday | day of the week
workingday | if day is neither weekend nor holiday

but it does not tell us what values are used in these columns.

Using `group`, we can see what is in `workingday` and how often then occur.

```python
bike.group("workingday")
```

Since most days are work days in real life, we can deduce that 1 means it was a workday and 0 means it was not.

**Question 1.3:** 
Assign `means_holiday` to the value that means that day was a holiday.
(Hint: Use group to figure out which values occur and their frequency.)

```python
...
means_holiday = ...
means_holiday
```

```python
grader.check("q1_3")
```

**Question 1.4:**
How many entries in the data correspond to holidays?  Set the variable `num_holidays` to this value.


<!-- grade_id: q1b-answer, points: 0, grade: false, solution: true -->
```python
num_holidays = ...
num_holidays
```

```python
grader.check("q1_4")
```

Similarly, we can use `group` to look at `weekday`.

```python
bike.group("weekday")
```

This tells us that the days are numbered 0 through 6, but it does not tell us if 0 means Sunday or Monday (or something else).

**Question 1.5:** 
Assign `monday_code` to the value used for Monday.
(Hint: Use group to figure out if 0 is a work day.)

```python
...
monday_code = ...
monday_code
```

```python
grader.check("q1_5")
```

**Question 1.6:** (Computing Daily Total Counts)
Construct a table named `daily_counts` with the following columns:
* `dteday` : as in `bike`, with each date occuring once
* `workingday` : as in `bike`
* `casual`: total number of casual riders for each day
* `registered`: total number of registered riders for each day


```python
daily_counts = ...
daily_counts
```

```python
grader.check("q1_6")
```

## 2: Exploring the Distribution of Riders

Some of the visualiation methods we use are available as methods of `datascience.Table`; see the [datascience documentation](http://www.data8.org/datascience/reference-nb/datascience-reference.html#Table-Visualizations).
We will also use `matplotlib` and `seaborn`.

### [Matplotlib](https://matplotlib.org/) Common Functions

See also the [tutorials](https://matplotlib.org/stable/tutorials/index.html)

`x` and `y` are sequences of values (i.e. arrays).

| Function | Description |
| -------- | ----------- |
| `plt.plot(x, y)` | Creates a line plot of `x` against `y` |
| `plt.title(name)` | Adds a title `name` to the current plot |
| `plt.xlabel(name)` | Adds a label `name` to the x-axis |
| `plt.ylabel(name)` | Adds a label `name` to the y-axis |
| `plt.legend()` | Add a legend to the plot, using labels set earlier with `label=...`|
| `plt.scatter(x, y)` | Creates a scatter plot of `x` against `y` |
| `plt.hist(x, bins=None)` | Creates a histogram of `x`; `bins` can be an integer or a sequence |
| `plt.bar(x, height)` | Creates a bar plot of categories `x` and corresponding heights `height` |

Generally, we will use another library to do the initial plot, and then adjust using `plt.title`, `plt.xlabel`, `plt.ylabel`, and `plt.legend()`. Here is an example of their usage. 

```python
plt.plot([0,1,2],[2,5,-1],label="my label for the legend")
plt.scatter([0,2,3],[0,1,2],label="my other label for the legend")
plt.title("My Title")
plt.xlabel("My x-axis label")
plt.ylabel("My y-axis label")
plt.legend()
```

### [Seaborn](https://seaborn.pydata.org/index.html) Table of Common Functions

See also the [tutorial](https://seaborn.pydata.org/tutorial.html)

`x` and `y` are sequences of values (i.e. arrays).
`data` is a `pandas.DataFrame`; if you have a `datascience.Table` called `tbl`, you would put `tbl.to_df()`.

| Function | Description |
| -------- | ----------- |
| `sns.histplot(data, x, y, hue, kde)` | Creates a distribution plot; `data` is a DataFrame; `x`, `y` are column names in `data` that specify positions on the x and y axes; `hue` is a column name in `data` that adds subcategories to the plot based on `hue`; `kde` is a boolean that determines whether to overlay a KDE curve |
|`sns.lineplot(data, x, y, hue)` | Creates a line plot |
|`sns.scatterplot(data, x, y, hue, size)` | Creates a scatter plot; `size` is a vector that contains the size of point for each subcategory based on `hue` |
|`sns.kdeplot(x, y)` |  Creates a kernel density estimate plot; `x`, `y` are series of data that indicate positions on the `x` and `y` axis |
|`sns.jointplot(x, y, data, kind)` | Creates a joint plot of 2 variables with KDE plot in the middle and a distribution plot for each variable on the sides; `kind` determines the visualization type for the distribution plot, can be `scatter`, `kde` or `hist` |

**Note**: This list of functions and parameters is **not** exhaustive. You may need to reference and explore more documentation to answer the following questions, but we will help you through that process. 

<!-- BEGIN QUESTION -->

**Question 2.1:**
Use the [`sns.histplot`](https://seaborn.pydata.org/generated/seaborn.histplot.html#seaborn.histplot) function to create a plot that overlays the distribution of the daily counts of bike users, using blue to represent `casual` riders, and green to represent `registered` riders. The temporal granularity of the records should be daily counts, so use `daily_counts`.

**Hint:** You will need to set the `stat`, `kde`, `label`, and `color` parameters appropriately to match the desired plot.

Include a legend, xlabel, ylabel, and title. Read the [seaborn plotting tutorial](https://seaborn.pydata.org/tutorial/distributions.html) if you're not sure how to add these. After creating the plot, look at it and make sure you understand what the plot is actually telling us, e.g on a given day, the most likely number of registered riders we expect is ~4000, but it could be anywhere from nearly 0 to 7000.

This is what your plot should look like.

<img src='casual_v_registered.png' width="600px" />


<!-- grade_id: q2a-answer, points: 2, grade: true, solution: true -->
```python
...
```

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

**Question 2.2:**
In the cell below, descibe the differences you notice between the density curves for casual and registered riders.  Consider concepts such as modes, symmetry, skewness, tails, gaps and outliers.  Include a comment on the spread of the distributions. 

_Type your answer here, replacing this text._

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

**Question 2.3:**
The density plots do not show us how the counts for registered and casual riders vary together. Use [`sns.lmplot`](https://seaborn.pydata.org/generated/seaborn.lmplot.html) to make a scatter plot to investigate the relationship between casual and registered counts. This time, let's use `bike` to plot hourly counts instead of daily counts.

The `lmplot` function will also try to draw a linear regression line. Color the points in the scatterplot according to whether or not the day is a working day (your colors do not have to match ours exactly, but they should be different based on whether the day is a working day). 

There are many points in the scatter plot, so make them small to help reduce overplotting. Also make sure to set `fit_reg=True` to generate the linear regression line. You can set the `height` parameter if you want to adjust the size of the `lmplot`. 

This is what your plot should look like:

<img src='casual_registered_working_nonworking.png' width="600px" />

**Hints:** 
* Checkout this helpful [tutorial on `lmplot`](https://seaborn.pydata.org/tutorial/regression.html).

* In the `sns.lmplot` call, you will need to set 
  * `x`,
  * `y`, 
  * `hue`, and
  * `scatter_kws={"s",point_size}` for some number `point_size` .
* You will need to call [`plt.title`](https://matplotlib.org/3.1.1/api/_as_gen/matplotlib.pyplot.title.html) to add a title for the graph.


<!-- grade_id: q2c-answer, points: 2, grade: true, solution: true -->
```python
# Make the font size a bit bigger
sns.set(font_scale=1)
...
```

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

**Question 2.4:**
What does this scatterplot seem to reveal about the relationship (if any) between casual and registered riders and whether or not the day is on the weekend? What effect does overplotting have on your ability to describe this relationship?

_Type your answer here, replacing this text._

<!-- END QUESTION -->

## 3: Bivariate Kernel Density Plots

To address overplotting, let's try visualizing the data with another technique, the bivariate kernel density estimate.
The result we wish to achieve is a plot that looks like this:

<img src='bivariate_kde_of_daily_rider_types.png' width="600px" />
<!-- plt.savefig("bivariate_kde_of_daily_rider_types.png", bbox_inches='tight', dpi=300);-->

A basic kde plot of all the data is quite easy to generate. However, this plot includes both weekend and weekday data, which isn't what we want (see example figure above).

```python
sns.kdeplot(x=daily_counts.column('casual'), y=daily_counts.column('registered'))
plt.title('KDE Plot Comparison of Registered vs Casual Riders');
```

<!-- BEGIN QUESTION -->

**Question 3.1:** Reproduce the plot above (with separate working days and non-working days).

Generating the plot with weekend and weekday separated can be complicated so we will provide a walkthrough below, feel free to use whatever method you wish if you do not want to follow the walkthrough.

**Hints:** 
* You will want to read up on the documentation for [`sns.kdeplot`](https://seaborn.pydata.org/generated/seaborn.kdeplot.html).
* You will need to call kdeplot twice, each time drawing different data from the `daily_counts` table.
* Check out this [guide](https://matplotlib.org/stable/tutorials/introductory/usage.html) to see an example of how to create a legend. In particular, look at how the example in the guide makes use of the `label` argument in the call to `plt.plot()` and what the `plt.legend()` call does. This is a good exercise to learn how to use examples to get the look you want.
* You will want to set the `cmap` parameter of `kdeplot` to `"Reds"` and `"Blues"` (or whatever two contrasting colors you'd like), and also set the `label` parameter to address which type of day you want to plot. You are required for this question to use two sets of contrasting colors for your plots.

After you get your plot working, experiment by setting `fill=True` in `kdeplot` to see the difference between the shaded and unshaded version. Please submit your work with `fill=False`. 

```python
# Set the figure size for the plot
plt.figure(figsize=(12,8))

# Bivariate KDEs require two data inputs. 
# In this case, we will need the daily counts for casual and registered riders on workdays 
casual_workday = ...
registered_workday = ...

# Use sns.kdeplot on the two variables above to plot the bivariate KDE for weekday rides
...

# Repeat the same steps above but for rows corresponding to non-workingdays 
casual_non_workday = ...
registered_non_workday = ...

# Use sns.kdeplot on the two variables above to plot the bivariate KDE for non-workingday rides
...

# set the labels and legend 
...
...
...
```

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

**Question 3.2:**
In your own words, describe what the lines and the color shades of the lines signify about the data.

_Type your answer here, replacing this text._

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

**Question 3.3:**
What additional details can you identify from this contour plot that were difficult to determine from the scatter plot?


_Type your answer here, replacing this text._

<!-- END QUESTION -->

An alternative approach to visualizing the data is to construct a joint plot.
The main plot shows the contours of the kernel density estimate of daily counts for registered and casual riders plotted together, and the two "margin" plots (at the top and right of the figure) provide the univariate kernel density estimate of each of these variables. Note that this plot makes it harder see the linear relationships between casual and registered for the two different conditions (weekday vs. weekend).

<img src="joint_distribution_of_daily_rider_types.png" width="600px" />

<!-- BEGIN QUESTION -->

**Question 3.4:**
Reproduce the joint plot above.

**Hints**:
* The [seaborn plotting tutorial](https://seaborn.pydata.org/tutorial/distributions.html) has examples that may be helpful.
* Take a look at [`sns.jointplot`](https://seaborn.pydata.org/generated/seaborn.jointplot.html) and its `kind` parameter.
* `set_axis_labels` can be used to rename axes on the contour plot.

**Note**:
* At the end of the cell, we called `plt.suptitle` to set a custom location for the title.
* We also called `plt.subplots_adjust(top=0.9)` in case your title overlaps with your plot.


<!-- grade_id: q3c-answer, points: 2, grade: true, solution: true -->
```python
...
plt.suptitle("KDE Contours of Casual vs Registered Rider Count")
plt.subplots_adjust(top=0.9);
```

<!-- END QUESTION -->

## 4. Time Series Plots

`bike` has information on how the number of riders changes over time. 
The `instant` column is the row index. Since rows are chronological, this gives us a way to plot versus time.

```python
bike.scatter("instant","registered")
```

Although it looks like there is an upward trend, the data is severely overplotted so it is difficult to see what is happening.

**Question 4.1:** Assign `hourly_means` to a table with 3 columns:
* `hr` as in `bike`, with each hour occuring once
* `casual mean` with the mean number of casual users for that hour
* `registered mean` with the mean number of registered users for that hour

```python
hourly_means = ...
hourly_means
```

```python
grader.check("q4_1")
```

We can now get a sense of the daily pattern.

```python
hourly_means.scatter("hr","registered mean")
```

However, a scatter plot seems to be the wrong type and it would be nice to see the casual rides too.
Here is a nicer graph.

<img src="diurnal_bikes.png" width="600px"/>

<!-- BEGIN QUESTION -->

**Question 4.2:**
Reproduce the graph above.
(Note: The above plot was produced using `sns.lineplot`. 
You could instead use `hourly_means.plot` to make something similar.)

<!-- grade_id: q5a-answer, points: 2, grade: true, solution: true -->
```python
...
```

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

**Question 4.3:**
What can you observe from the plot?  Hypothesize about the meaning of the peaks in the registered riders' distribution.

_Type your answer here, replacing this text._

<!-- END QUESTION -->

Here is a plot produced while playing around with the data:

```python
bike.where('weekday',are.equal_to(3)).group('dteday',np.mean).select('instant mean','casual mean','registered mean').plot("instant mean")
```

<!-- BEGIN QUESTION -->

**Question 4.4:**
What is this plotting? 
What information does it give?

_Type your answer here, replacing this text._

<!-- END QUESTION -->

## Acknowledgements

<!-- BEGIN QUESTION -->

**Acknowledgements:**
If you received help from any person other than your instructor, then describe here what help you received and from whom. 
If not, then say 'None'.

_Type your answer here, replacing this text._

<!-- END QUESTION -->

It is forbidden to: 
* Use the work or programs from students who took this class (in any of its versions at any university).
* Use websites that claim to have solutions for this class (in any of its versions at any university).
* Directly copy from a classmate.

**Did you do anything forbidden while working on this assignment?**

Assign `did_forbidden` to either `True` or `False`.

```python
did_forbidden = ...
```

```python
grader.check("q_forbidden")
```

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit. **Please save before exporting!**

After exporting, submit the .zip file in BlackBoard

```python
# Save your notebook first, then run this cell to export your submission.
grader.export(run_tests=True)
```

Tests for q1_3 (points: None)
```
>>> means_holiday in [0, 1]
True
```


Tests for q1_4 (points: None)
```
>>> 1 <= num_holidays <= 10000
True
```


Tests for q1_5 (points: None)
```
>>> monday_code in [0, 1]
True
```


Tests for q1_6 (points: None)
```
>>> daily_counts.labels
('dteday', 'workingday', 'casual', 'registered')
```
```
>>> np.round(np.mean(daily_counts.column('casual')))
848.0
```
```
>>> np.round(np.var(daily_counts.column('casual')))
470806.0
```


Tests for q4_1 (points: None)
```
>>> hourly_means.num_rows
24
```
```
>>> len(hourly_means.labels)
3
```


Tests for q_forbidden (points: None)
```
>>> did_forbidden in [True, False]
True
```
