```python
# Initialize Otter
import otter
grader = otter.Notebook("lab07.ipynb")
```

# Lab 7: More Visualization, Transformations, and KDEs
(This lab is based on [Data 100 Fall 2022 lab 04](https://github.com/DS-100/fa22/tree/main/lab/lab04).)

### Objective
In this lab you will get some practice plotting, applying data transformations, and working with kernel density estimators (KDEs).  We will be working with data from the World Bank containing various statistics for countries and territories around the world. 

```python
#import pandas as pd
import numpy as np
import seaborn as sns
import matplotlib.pyplot as plt
from IPython.display import display, Markdown, Latex
from datascience import *

plt.style.use('fivethirtyeight') # Use plt.style.available to see more styles
sns.set()
sns.set_context("talk")
%matplotlib inline
```

## 1. Loading and Preparing Data

Let us load some World Bank data.

```python
wb_raw = Table.read_table('world_bank_misc.csv')
wb_raw.show(3)
```

We can see the names of all the columns.

```python
wb_raw.labels
```

From these, we are going to focus on income and adult literacy. 

**Question 1.1:** From `wb_raw`, create a table `wb_income_literacy` with 4 columns:
* `country`, with the data from the column `Unnamed: 0`
* `income`, with the data from the column `Gross national income per capita, Atlas method: $: 2016`
* `male literacy`, with the data from the column `Adult literacy rate: Male: % ages 15 and older: 2005-14`
* `female literacy`, with the data from the column `Adult literacy rate: Female: % ages 15 and older: 2005-14`

```python
wb_income_literacy = ...
                                          ...
                                          ...
                                          ...
                                         )
wb_income_literacy
```

```python
grader.check("q1_1")
```

Notice that some countries have missing data, marked with `nan`. 
All the other values are positive, so we can exclude `nan` using the predicate `are.above(0)`.

**Question 1.2:** 
From `wb_income_literacy`, create a table `wb_income_literacy_nonan` that skips rows that have a `nan` in any column.

```python
wb_income_literacy_nonan = ...
wb_income_literacy_nonan
```

```python
grader.check("q1_2")
```

**Question 1.3:** Assign `wb` to the table that has everything from `wb_income_literacy_nonan`, plus a column `literacy` with the mean of the `male literacy` and the `female literacy`. 

```python
literacy_array = ...
wb = ...
wb
```

```python
grader.check("q1_3")
```

## 2. Various way to plot distributions in one variable

In order to understand them better, we would like to have visualizations of our data to understand the distribution of literacy rates and income per capita individually. 

Some of these visualizations uses methods built in to the `Table` type.
Others are from the [`seaborn`](https://seaborn.pydata.org/) library, which expects `pandas.DataFrame`s as inputs instead of `datascience.Table`s. To use these, we will need to convert our `Table` using the `.to_df()` method.


### 2.1 Boxplots

A [boxplot](https://en.wikipedia.org/wiki/Box_plot) displays
* the median, as a line
* the interquartile range, as a box
* the min and max, as whiskers, but with points thought to be outliers as separate points

This is "traditional" type of plot, which can be drawn with limited (or no) computer resources.
It can be produced [directly from a `datascience` `Table`](http://www.data8.org/datascience/_autosummary/datascience.tables.Table.boxplot.html#datascience.tables.Table.boxplot).

Here is how to produce the boxplot for `income`.

```python
wb.select('income').boxplot()
```

<!-- BEGIN QUESTION -->

**Question 2.1.1:** Produce a single plot that has the boxplots for `male literacy`, `female literacy`, and `literacy` side-by-side.

```python
...
```

<!-- END QUESTION -->

### 2.2 Countplot

We can use [`countplot`](https://seaborn.pydata.org/generated/seaborn.countplot.html) in seaborn to create bar charts from categorical data. 

```python
sns.countplot(x = "literacy", data = wb.to_df())
plt.xlabel("Combined literacy rate: % ages 15 and older: 2005-14")
plt.title('World Bank Combined Adult Literacy Rate')
```

```python
sns.countplot(x = "income", data = wb.to_df())
plt.xlabel('Gross national income per capita, Atlas method: $: 2016')
plt.title('World Bank Gross National Income Per Capita')
```

<!-- BEGIN QUESTION -->

**Question 2.2.1:** In the cell below, explain why `countplot` is NOT the right tool for visualizing the distribution of our data.


_Type your answer here, replacing this text._

<!-- END QUESTION -->

### 2.3 Histograms

<!-- BEGIN QUESTION -->

**Question 2.3.1:**
Create a plot of **income per capita** (the second plot above) using the [`histplot`](https://seaborn.pydata.org/generated/seaborn.histplot.html) function. 

Don't forget to title the plot and label axes!

**Hint:** *Copy and paste from above to start.*


<!-- grade_id: q2a-answer, points: 1, grade: true, solution: true -->
```python
...
```

<!-- END QUESTION -->

You should see a histogram that show the counts of how many data points appear in each bin. `histplot` uses a heuristic called the Freedman-Diaconis rule to automatically identify the best bin sizes, though it is possible to set the bins yourself (we won't).

### 2.4 Rugplot

In the cell below, we explore overlaying a rug plot on top of a histogram using `rugplot`. Note that the rug plot is hard to see.

<!-- grade_id: cell-72d2c195130debd8, points: 1, grade: true, solution: true -->
```python
sns.histplot(x="income", data = wb.to_df())
sns.rugplot(x="income", data = wb.to_df())
plt.xlabel('Gross national income per capita, Atlas method: $: 2016')
plt.title('World Bank Gross National Income Per Capita')
```

One way to make it easier to see the difference between the rug plot and the bars is to set a different color, for example:

```python
sns.histplot(x="income", data = wb.to_df(), color = "lightsteelblue")
sns.rugplot(x="income", data = wb.to_df())
plt.xlabel('Gross national income per capita, Atlas method: $: 2016')
plt.title('World Bank Gross National Income Per Capita')
```

### 2.5 Kdeplot

A [Kernel Density Estimation](https://en.wikipedia.org/wiki/Kernel_density_estimation) gives a sort of smoothed-out histogram. 
Later in this lab we will consider how it gets this curve.

The seaborn function  [`kdeplot`](https://seaborn.pydata.org/generated/seaborn.kdeplot.html) produces this type of plot. 
Here is how to make the kdeplot for `income`.

```python
ax = sns.kdeplot(x="income", data = wb.to_df())
```

### 2.6 Violin Plots

A [violin plot](https://en.wikipedia.org/wiki/Violin_plot) has aspects of both a boxplot and a kernel density estimation.

We can make them using seaborn's [`violinplot`](https://seaborn.pydata.org/generated/seaborn.violinplot.html#seaborn.violinplot).

Here is how to produce a violinplot for `income`. Note that the kernel density estimate gives some negative incomes.

```python
ax = sns.violinplot(data = wb.select('income').to_df())
```

<!-- BEGIN QUESTION -->

**Question 2.6.1:** Produce a single plot that has the violinplots for `male literacy`, `female literacy`, and `literacy` side-by-side.

```python
ax = ...
```

<!-- END QUESTION -->

### 2.7 Displot

Rather than manually calling `histplot`, `rugplot`, and `kdeplot` to plot histograms, rug plots, and KDE plots, respectively, we can instead use `displot`, which can simultaneously plot histogram bars, a rug plot, and a KDE plot, and adjust all the colors automatically for visbility. Using the documentation for [`displot`](https://seaborn.pydata.org/generated/seaborn.displot.html), make a plot of the income data that includes a histogram, rug plot, and KDE plot. Hint: You'll need to set two parameters to `True`.

<!-- grade_id: cell-1aaac759f52a2269, points: 1, grade: true, solution: true -->
```python
ax = sns.displot(x="income", data = wb.to_df(), rug = True, kde = True)
plt.xlabel('Gross national income per capita, Atlas method: $: 2016')
plt.title('World Bank Gross National Income Per Capita')
```

You should see roughly the same histogram as before. However, now you should see an overlaid smooth line. This is the kernel density estimate. 

Above, the y-axis is labeled by the counts. We can also label the y-axis by the density. An example is given below, this time using the literacy data from the beginning of this lab.

```python
ax = sns.displot(x="literacy", data = wb.to_df(), rug = True, kde = True, stat = "density")
plt.xlabel("Adult literacy rate: Combined: % ages 15 and older: 2005-14")
plt.title('World Bank Combined Adult Literacy Rate')
```

Observations:
* You'll also see that the y-axis value is no longer the count. Instead it is a value such that the total **area** in the histogram is 1. For example, the area of the last bar is approximately 22.22 * 0.028 = 0.62

* The KDE is a smooth estimate of the distribution of the given variable. The area under the KDE is also 1. While it is not obvious from the figure, some of the area under the KDE is beyond the 100% literacy. In other words, the KDE is non-zero for values greater than 100%. This, of course, makes no physical sense. Nonetheless, it is a mathematical feature of the KDE.

We'll talk more about KDEs later in this lab.

### 2.8 Transforming data to improve visualizations

Looking at the income data, it is difficult to see the distribution among low income countries because they are all scrunched up at the left side of the plot. The KDE also has a problem where the density function has a lot of area below 0. 

Transforming the `income` data logarithmically gives us a more symmetric distribution of values. This can make it easier to see patterns.

In addition, summary statistics like the mean and standard deviation (square-root of the variance) are more stable with symmetric distributions.

<!-- BEGIN QUESTION -->

**Question 2.8.1:**
Make a distribution plot of `income` with the data transformed using `np.log10` and `kde=True`. If you want to see the exact counts, just set `kde=False`. If you don't specify the `kde` parameter, it is by default set to True. 

**Hint:** Unlike the examples above, you can pass an array to the `displot` function, i.e. rather than passing an entire Table as `data` and a column as `x`, you can instead pass the array that you made from a column.


<!-- grade_id: q2c-answer, points: 1, grade: true, solution: true -->
```python
ax = ...
plt.title('World Bank Gross National Income Per Capita')
plt.ylabel('Density')
plt.xlabel('Log Gross national income per capita, Atlas method: $: 2016');
```

<!-- END QUESTION -->

When a distribution has a long right tail, a log-transformation often does a good job of symmetrizing the distribution, as it did here.  Long right tails are common with variables that have a lower limit on the values. 

On the other hand, long left tails are common with distributions of variables that have an upper limit, such as percentages (can't be higher than 100%) and GPAs (can't be higher than 4).  That is the case for the literacy rate. Typically taking a power-transformation such 
as squaring or cubing the values can help symmetrize the left skew distribution.

In the cell below, we will make a distribution plot of `literacy` with the data transformed using a power, i.e., raise `literacy` to the 2nd, 3rd, and 4th power. We plot the transformation with the 4th power below.


```python
ax = sns.displot((wb.column('literacy')**4), kde = True)
plt.ylabel('Density')
plt.xlabel("Adult literacy rate: Combined: % ages 15 and older: 2005-14")
plt.title('World Bank Combined Adult Literacy Rate (4th power)', pad=30);
```

## 3. Plotting relationships between variables

If we want to examine the relationship between the combined adult literacy rate and the gross national income per capita, we need to make a scatter plot. 


<!-- BEGIN QUESTION -->

**Question 3.1:**
In the cell below, create a scatter plot of untransformed income per capita and literacy rate using the `sns.scatterplot` function. Make  sure to label both axes using `plt.xlabel` and `plt.ylabel`.


```python
...
```

<!-- END QUESTION -->

We can better assess the relationship between two variables when they have been straightened because it is easier for us to recognize linearity.

In the cell below, we see a scatter plot of log-transformed income per capita against literacy rate.


```python
sns.scatterplot(x = wb.column('literacy'), y = wb.apply(np.log10,'income'))
plt.xlabel("Adult literacy rate: Combined: % ages 15 and older")
plt.ylabel('Gross national income per capita (log scale)')
plt.title('World Bank: Gross National Income Per Capita vs\n Combined Adult Literacy Rate');
```

This scatter plot looks better. The relationship is closer to linear.

We can think of the log-linear relationship between x and y, as follows: a constant change in x corresponds to a percent (scaled) change in y.

We can also see that the long left tail of literacy is represented in this plot by a lot of the points being bunched up near 100. 

<!-- BEGIN QUESTION -->

**Question 3.2:** Try squaring literacy and taking the log of income. Does the plot look better? 


```python
plt.figure(figsize=(10,5))
...
```

<!-- END QUESTION -->

Choosing the best transformation for a relationship is often a balance between keeping the model simple and straightening the scatter plot.

## 4. Kernel Density Estimation

In this part of the lab you will develop a deeper understanding of how kernel density estimation works.

### Overview

Kernel density estimation is used to estimate a probability density function (i.e. a density curve) from a set of data. Just like a histogram, a density function's total area must sum to 1.

KDE centrally revolves around this idea of a "kernel". A kernel is a function whose area sums to 1. The three steps involved in building a kernel density estimate are:
1. Placing a kernel at each observation
2. Normalizing kernels so that the sum of their areas is 1
3. Summing all kernels together

The end result is a function, that takes in some value `x` and returns a density estimate at the point `x`.

When constructing a KDE, there are several choices to make regarding the kernel. Specifically, we need to choose the function we want to use as our kernel, as well as a bandwidth parameter, which tells us how wide or narrow each kernel should be. We will explore these ideas now.

Suppose we have 3 data points with values 2, 4, and 9. We can compute the (useless) histogram with a KDE as shown below.

```python
data3pts = np.array([2, 4, 9])
sns.displot(data3pts, kde = True, stat = "density");
```

To understand how KDEs are computed, we need to see the KDE outside the given range. The easiest way to do this is to use an old function called `distplot`. During the Spring 2023 offering of this course, `distplot` was still a working function in Seaborn, but it will be removed at a future date. If you get an error that says that `distplot` is not a valid function, sorry, you are too far in the future to do this lab exercise.

```python
sns.distplot(data3pts, kde = True);
```

One question you might be wondering is how the kernel density estimator decides how "wide" each point should be. It turns out this is a parameter you can set called `bw`, which stands for bandwith. For example, the code below gives a bandwith value of 0.5 to each data point. You'll see the resulting KDE is quite different. Try experimenting with different values of bandwidth and see what happens.

```python
sns.distplot(data3pts, kde = True, kde_kws = {"bw": 0.5});
```

As mentioned above, the kernel density estimate (KDE) is just the sum of a bunch of copies of the kernel, each centered on our data points. The default kernel used by the `distplot` function (as well as `kdeplot`) is the Gaussian kernel, given by:

$$\Large
K_\alpha(x, z) = \frac{1}{\sqrt{2 \pi \alpha^2}} \exp\left(-\frac{(x - z)^2}{2  \alpha ^2} \right)
$$

We've implemented the Gaussian kernel for you in Python below. Here, `alpha` is the smoothing or bandwidth parameter $\alpha$ for the KDE, `z` is the center of the Gaussian (i.e. a data point or an array of data points), and `x` is an array of values of the variable whose distribution we are plotting.

```python
def gaussian_kernel(alpha, x, z):
    return 1.0/np.sqrt(2. * np.pi * alpha**2) * np.exp(-(x - z) ** 2 / (2.0 * alpha**2))
```

For example, we can plot the Gaussian kernel centered at 9 with $\alpha$ = 0.5 as below: 

```python
xs = np.linspace(-2, 12, 200)
alpha=0.5
kde_curve = [gaussian_kernel(alpha, x, 9) for x in xs]
plt.plot(xs, kde_curve);
```

Recall our original (useless) plot of our 3 data points:

```python
sns.displot(data3pts, kde = True, stat = "density");
```

We can now plot the density from our 3 data points while varying the width. Try different values of alpha. 

<!-- grade_id: cell-a1cbcf4340aacbdd, points: 1, grade: true, solution: true -->
```python
xs = np.linspace(2, 9, 200)
alpha=0.2
kde_curve = np.array([1/len(data3pts) * gaussian_kernel(alpha, x, data3pts) for x in xs])
plt.plot(xs, np.sum(kde_curve, axis = 1));
plt.ylim([0,plt.ylim()[1]])
```

<!-- BEGIN QUESTION -->

**Question 4.1:**
Which value of `alpha` best matches the curve in our original plot?

_Type your answer here, replacing this text._

<!-- END QUESTION -->

### KDE on the income data
Recall that earlier we plotted the kernel density estimation for the logarithm of the income data, as shown again below.

```python
ax = sns.displot(np.log10(wb.column('income')), kind = "kde", rug = True)
plt.title('World Bank Gross National Income Per Capita')
plt.xlabel('Log Gross national income per capita, Atlas method: $: 2016');
```

Try out different values of alpha in {0.1, 0.2, 0.3, 0.4, 0.5}. You will see that when alpha=0.2, the graph matches the previous graph well, except that the `displot` function hides the KDE values outside the range of the available data.

<!-- grade_id: cell-a1cbcf4340aacbdd, points: 1, grade: true, solution: true -->
```python
xs = np.linspace(1, 6, 200)
alpha=0.2
kde_curve = np.array([1/(wb.num_rows) * gaussian_kernel(alpha, x, np.log10(wb.column('income'))) for x in xs])
plt.title('World Bank Gross National Income Per Capita')
plt.xlabel('Log Gross national income per capita, Atlas method: $: 2016')
plt.plot(xs, np.sum(kde_curve, axis = 1));
```

The code below generates the `kde` of the log of the income data as before.

```python
transformed_income = np.log10(wb.column('income'))
alpha = 0.2
kde = lambda a,b,c,d : np.sum(a(b,c,d))/len(d)
xs = np.linspace(min(transformed_income), max(transformed_income), 1000)
curve = [kde(gaussian_kernel, alpha, x, transformed_income) for x in xs]
plt.hist(transformed_income, density=True, color='orange')
plt.title('World Bank Gross National Income Per Capita')
plt.xlabel('Log Gross national income per capita, Atlas method: $: 2016');
plt.plot(xs, curve, 'k-');
```

And the code below should show a 3 x 3 set of plots showing the output of the kde for different `alpha` values.

```python
plt.figure(figsize=(15,15))
alphas = np.arange(0.2, 2.0, 0.2)
kde = lambda a,b,c,d : np.sum(a(b,c,d))/len(d)
xs = np.linspace(min(transformed_income), max(transformed_income), 1000)
for i, alpha in enumerate(alphas):
    plt.subplot(3, 3, i+1)   
    curve = [kde(gaussian_kernel, alpha, x, transformed_income) for x in xs]
    plt.hist(transformed_income, density=True, color='orange')
    plt.plot(xs, curve, 'k-')
    plt.title("alpha="+str(alpha))
plt.show()
```

### The Boxcar Kernel
Let's take a look at another kernel, the Boxcar kernel.

```python
def boxcar_kernel(alpha, x, z):
    return (((x-z)>=-alpha/2)&((x-z)<=alpha/2))/alpha
```

Run the cell below to enable interactive plots. It should give you a green 'OK' when it's finished.

```python
from ipywidgets import interact
!jupyter nbextension enable --py widgetsnbextension
```

Now, we can plot the Boxcar and Gaussian kernel functions to see what they look like.

```python
x = np.linspace(-10,10,1000)
def f(alpha):
    plt.plot(x, boxcar_kernel(alpha,x,0), label='Boxcar')
    plt.plot(x, gaussian_kernel(alpha,x,0), label='Gaussian')
    plt.legend(title='Kernel Function')
    plt.show()
interact(f, alpha=(1,10,0.1));
```

Using the interactive plot below compare the the two kernel techniques:  (Generating the KDE plot is slow, so you may expect some latency after you move the slider)

```python
xs = np.linspace(min(transformed_income), max(transformed_income), 1000)
kde = lambda a,b,c,d : np.sum(a(b,c,d))/len(d)
def f(alpha_g, alpha_b):
    plt.hist(transformed_income, density=True, color='orange')
    g_curve = [kde(gaussian_kernel, alpha_g, x, transformed_income) for x in xs]
    plt.plot(xs, g_curve, 'k-', label='Gaussian')
    b_curve = [kde(boxcar_kernel, alpha_b, x, transformed_income) for x in xs]
    plt.plot(xs, b_curve, 'r-', label='Boxcar')
    plt.legend(title='Kernel Function')
    plt.show()
interact(f, alpha_g=(0.01,.5,0.01), alpha_b=(0.01,3,0.1));
```

<!-- BEGIN QUESTION -->

**Question 4.2:** Briefly compare and contrast the Gaussian and Boxcar kernels in the cell below. How do the two kernels relate with each other for the same alpha value?

_Type your answer here, replacing this text._

<!-- END QUESTION -->

**Congrats!** You are finished with this lab.

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit. **Please save before exporting!**

After exporting, submit the .zip file in BlackBoard

```python
# Save your notebook first, then run this cell to export your submission.
grader.export(run_tests=True)
```

Tests for q1_1 (points: None)
```
>>> wb_income_literacy.num_rows
216
```
```
>>> set(wb_income_literacy.labels) == set(['country', 'income', 'male literacy', 'female literacy'])
True
```
```
>>> wb_income_literacy.column('country').item(10)
'Australia'
```


Tests for q1_2 (points: None)
```
>>> wb_income_literacy_nonan.num_rows
147
```
```
>>> wb_income_literacy_nonan.column('country').item(10)
'Belarus'
```


Tests for q1_3 (points: None)
```
>>> wb.num_rows
147
```
```
>>> set(wb.labels) == set(['country', 'income', 'male literacy', 'female literacy', 'literacy'])
True
```
```
>>> 99.4 < wb.column('literacy').item(10) < 99.6
True
```
