```python
# Initialize Otter
import otter
grader = otter.Notebook("hw05.ipynb")
```

# Homework 5: Visualization

<!-- BEGIN QUESTION -->

**Question:** What is your name?

_Type your answer here, replacing this text._

<!-- END QUESTION -->

Please complete this notebook by filling in the cells provided. 

**Throughout this homework and all future ones, please be sure to not re-assign variables throughout the notebook!** For example, if you use `max_temperature` in your answer to one question, do not reassign it later on. Moreover, please be sure to only put your written answers in the provided cells. 

```python
# Don't change this cell; just run it. 

import numpy as np
from datascience import *


# These lines do some fancy plotting magic.\n",
import matplotlib
%matplotlib inline
import matplotlib.pyplot as plt
plt.style.use('fivethirtyeight')


```

## 1. Unemployment
 (This section is based on [Data 8 hw 03](https://github.com/data-8/materials-sp20/tree/master/materials/sp20/hw/hw03) section 1.)

The Federal Reserve Bank of St. Louis publishes data about jobs in the US.  Below, we've loaded data on unemployment in the United States. There are many ways of defining unemployment, and our dataset includes two notions of the unemployment rate:

1. Among people who are able to work and are looking for a full-time job, the percentage who can't find a job.  This is called the Non-Employment Index, or NEI.
2. Among people who are able to work and are looking for a full-time job, the percentage who can't find any job *or* are only working at a part-time job.  The latter group is called "Part-Time for Economic Reasons", so the acronym for this index is NEI-PTER.  (Economists are great at marketing.)

The source of the data is [here](https://fred.stlouisfed.org/categories/33509).

**Question 1.** The data are in a CSV file called `unemployment.csv`.  Load that file into a table called `unemployment`.


```python
unemployment = ...
unemployment
```

```python
grader.check("q1_1")
```

**Question 2.** It's believed that many people became PTER (recall: "Part-Time for Economic Reasons") in the "Great Recession" of 2008-2009.  NEI-PTER is the percentage of people who are unemployed (and counted in the NEI) plus the percentage of people who are PTER.  Compute an array containing the percentage of people who were PTER in each quarter.  (The first element of the array should correspond to the first row of `unemployment`, and so on.)

```python
pter = ...
pter
```

```python
grader.check("q1_2was1_4")
```

**Question 3.** 

Create a line plot of the PTER over time. 

To do this, create a new table called `pter_over_time` that adds the `year` array and the `pter` array to the `unemployment` table. Label these columns `Year` and `PTER`. Then, generate a line plot using one of the table methods you've learned in class.


```python
year = 1994 + np.arange(len(pter))/4
pter_over_time = ...
...
```

```python
grader.check("q1_3was1_6")
```

**Question 4.** Were PTER rates high during the Great Recession (that is to say, were PTER rates particularly high in the years 2008 through 2011)? Assign highPTER to `True` if you think PTER rates were high in this period, and `False` if you think they weren't. 


```python
highPTER = ...
```

```python
grader.check("q1_4was1_7")
```

## 2. Birth Rates
 (This section is based on [Data 8 hw 03](https://github.com/data-8/materials-sp20/tree/master/materials/sp20/hw/hw03) section 2.)

The following table gives census-based population estimates for each state on both July 1, 2015 and July 1, 2016. The last four columns describe the components of the estimated change in population during this time interval. **For all questions below, assume that the word "states" refers to all 52 rows including Puerto Rico & the District of Columbia.**

The data was taken from [here](http://www2.census.gov/programs-surveys/popest/datasets/2010-2016/national/totals/nst-est2016-alldata.csv).

If you want to read more about the different column descriptions, click [here](http://www2.census.gov/programs-surveys/popest/datasets/2010-2015/national/totals/nst-est2015-alldata.pdf)!

The raw data is a bit messy - run the cell below to clean the table and make it easier to work with.

```python
# Don't change this cell; just run it.
pop = Table.read_table('nst-est2016-alldata.csv').where('SUMLEV', 40).select([1, 4, 12, 13, 27, 34, 62, 69])
pop = pop.relabeled('POPESTIMATE2015', '2015').relabeled('POPESTIMATE2016', '2016')
pop = pop.relabeled('BIRTHS2016', 'BIRTHS').relabeled('DEATHS2016', 'DEATHS')
pop = pop.relabeled('NETMIG2016', 'MIGRATION').relabeled('RESIDUAL2016', 'OTHER')
pop = pop.with_columns("REGION", np.array([int(region) if region != "X" else 0 for region in pop.column("REGION")]))
pop.set_format([2, 3, 4, 5, 6, 7], NumberFormatter(decimals=0)).show(5)
```

**Question 1.** 

In the next question, you will be creating a visualization to understand the relationship between birth and death rates. The annual death rate for a year-long period is the total number of deaths in that period as a proportion of the population size at the start of the time period.

What visualization is most appropriate to see if there is an association between birth and death rates during a given time interval?

1. Line Graph
<br>
2. Scatter Plot
<br>
3. Bar Chart

Assign `visualization` below to the number corresponding to the correct visualization.


```python
visualization = ...
```

```python
grader.check("q2_1was2_5")
```

<!-- BEGIN QUESTION -->

**Question 2.** 

In the code cell below, create a visualization that will help us determine if there is an association between birth rate and death rate during this time interval. It may be helpful to create an intermediate table here. 
Use one of the `Table` methods `.plot`, `.scatter`, `.bar`, or `.barh`.


```python
# Generate your chart in this cell
...
```

<!-- END QUESTION -->

**Question 3.** `True` or `False`: There is an association between birth rate and death rate during this time interval. 

Assign `assoc` to `True` or `False` in the cell below. 


```python
assoc = ...
```

```python
grader.check("q2_3")
```

## 3. Marginal Histograms
(This section is based on [Data 8 hw 03](https://github.com/data-8/materials-sp20/tree/master/materials/sp20/hw/hw03) section 3.)

Consider the following scatter plot: ![](scatter.png)

The axes of the plot represent values of two variables: $x$ and $y$. 

Suppose we have a table called `t` that has two columns in it:

- `x`: a column containing the x-values of the points in the scatter plot
- `y`: a column containing the y-values of the points in the scatter plot

Below, you are given two histograms, each of which corresponds to either column `x` or column `y`. 

**Histogram A:** ![](var1.png)
**Histogram B:** ![](var2.png)

**Question 1.** Suppose we run `t.hist('x')`. Which histogram does this code produce? Assign `histogram_column_x` to either 1 or 2.

1. Histogram A
2. Histogram B


```python
histogram_column_x = ...
```

```python
grader.check("q3_1")
```

<!-- BEGIN QUESTION -->

**Question 2.** State at least one reason why you chose the histogram from Question 1. Make sure to indicate which histogram you selected (ex: "I chose histogram A because ...").


_Type your answer here, replacing this text._

<!-- END QUESTION -->

**Question 3.** Suppose we run `t.hist('y')`. Which histogram does this code produce? `Assign histogram_column_y` to either 1 or 2.

1. Histogram A
2. Histogram B


```python
histogram_column_y = ...
```

```python
grader.check("q3_3")
```

<!-- BEGIN QUESTION -->

**Question 4.** State at least one reason why you chose the histogram from Question 3.  Make sure to indicate which histogram you selected (ex: "I chose histogram A because ...").


_Type your answer here, replacing this text._

<!-- END QUESTION -->

## 4. Uber
(This section is based on [Data 8 hw 03](https://github.com/data-8/materials-sp20/tree/master/materials/sp20/hw/hw03) section 4.)

Below we load tables containing 200,000 weekday Uber rides in the Manila, Philippines, and Boston, Massachusetts metropolitan areas from the [Uber Movement](https://movement.uber.com) project. The `sourceid` and `dstid` columns contain codes corresponding to start and end locations of each ride. The `hod` column contains codes corresponding to the hour of the day the ride took place. The `ride time` column contains the length of the ride, in minutes.

```python
boston = Table.read_table("boston.csv")
manila = Table.read_table("manila.csv")
print("Boston Table")
boston.show(4)
print("Manila Table")
manila.show(4)
```

<!-- BEGIN QUESTION -->

**Question 1.** Produce histograms of all ride times in Boston using the given bins. 



```python
equal_bins = np.arange(0, 120, 5)
...
```

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

**Question 2.** Now, produce histograms of all ride times in Manila using the given bins.


```python
...

# Don't delete the following line!
plt.ylim(0, 0.05)
```

<!-- END QUESTION -->

**Question 3.** Assign `boston_under_10` and `manila_under_10` to the percentage of rides that are less than 10 minutes in their respective metropolitan areas. Use the height variables provided below in order to compute the percentages. Your solution should only use height variables, numbers, and mathematical operations. You should not access the tables boston and manila in any way.


```python
boston_under_5_height = 1.2
manila_under_5_height = 0.6
boston_5_to_under_10_height = 3.2
manila_5_to_under_10_height = 1.4

boston_under_10 = ...
manila_under_10 = ...
```

```python
grader.check("q4_3")
```

**Question 4.** Let's take a closer look at the distribution of ride times in Manila. Assign `manila_median_bin` to an integer (1, 2, 3, or 4) that corresponds to the bin that contains the median time 

1: 0-15 minutes  
2: 15-40 minutes  
3: 40-60 minutes  
4: 60-80 minutes  

*Hint:* The median of a sorted list has half of the list elements to its left, and half to its right


```python
manila_median_bin = ...
manila_median_bin
```

```python
grader.check("q4_4")
```

<!-- BEGIN QUESTION -->

**Question 5.** What is the main difference between the two histograms. What might be causing this?

**Hint:** Try thinking about external factors that may be causing the difference!


_Type your answer here, replacing this text._

<!-- END QUESTION -->

### 5. Burrito-ful San Diego
(This section is based on [Data 8 hw 04](https://github.com/data-8/materials-sp20/tree/master/materials/sp20/hw/hw04) section 1.)

Tam, Margaret and Winifred are trying to use Data Science to find the best burritos in San Diego! Their friends Irene and Maya provided them with two comprehensive datasets on many burrito establishments in the San Diego area taken from (and cleaned from): https://www.kaggle.com/srcole/burritos-in-san-diego/data

The following cell reads in a table called `ratings` which contains names of burrito restaurants, their Yelp rating, Google rating, as well as their Overall rating. It also reads in a table called `burritos_types` which contains names of burrito restaurants, their menu items, and the cost of the respective menu item at the restaurant.

```python
#Just run this cell
ratings = Table.read_table("ratings.csv")
ratings.show(5)
burritos_types = Table.read_table("burritos_types.csv")
burritos_types.show(5)
```

**Question 1.** It would be easier if we could combine the information in both tables. Assign `burritos` to the result of joining the two tables together.

*Note: it doesn't matter which table you put in as the argument to the table method, either order will work for the autograder tests.*

*Hint: If you need refreshers on table methods, look at the [python reference](http://data8.org/sp20/python-reference.html).*


```python
burritos = ...
burritos.show(5)
```

```python
grader.check("q5_1was1_1")
```

<!-- BEGIN QUESTION -->

**Question 2.** Let's look at how the Yelp scores compare to the Google scores in the `burritos` table. First, assign `yelp_and_google` to a table only containing the columns `Yelp` and `Google`. Then, make a scatter plot with Yelp scores on the x-axis and the Google scores on the y-axis. 


```python
yelp_and_google = ...
...
# Don't change/edit/remove the following line.
# To help you make conclusions, we have plotted a straight line on the graph (y=x)
plt.plot(np.arange(2.5,5,.5), np.arange(2.5,5,.5));
```

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

**Question 3.** Looking at the scatter plot you just made in Question 1.2, do you notice any pattern(s) (i.e. is one of the two types of scores consistently higher than the other one)? If so, describe them **briefly** in the cell below.


_Type your answer here, replacing this text._

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

**Question 4.** Tam thinks that burritos in San Diego are cheaper (and taste better) than the burritos in Berkeley. Plot a histogram that visualizes that distribution of the costs of the burritos from San Diego in the `burritos` table. Also use the provided `bins` variable when making your histogram, so that visually the histogram is more informative.


```python
bins = np.arange(0, 15, 1)
# Please also use the provided bins
...
```

<!-- END QUESTION -->

## Acknowledgements

<!-- BEGIN QUESTION -->

**Acknowledgements:**
If you received help from any person other than your instructor, then describe here what help you received and from whom. 
If not, then say 'None'.

_Type your answer here, replacing this text._

<!-- END QUESTION -->

It is forbidden to: 
* Use the work or programs from students who took this class (in any of its versions at any university).
* Use websites that claim to have solutions for this class (in any of its versions at any university).
* Directly copy from a classmate.

**Did you do anything forbidden while working on this assignment?**

Assign `did_forbidden` to either `True` or `False`.

```python
did_forbidden = ...
```

```python
grader.check("q_forbidden")
```

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit. **Please save before exporting!**

After exporting, submit the .zip file in BlackBoard

```python
# Save your notebook first, then run this cell to export your submission.
grader.export(run_tests=True)
```

Tests for q1_1 (points: None)
```
>>> unemployment.select('Date', 'NEI', 'NEI-PTER').take(0).column(0).item(0) == '1994-01-01'
True
```


Tests for q1_2was1_4 (points: None)
```
>>> round(pter.item(6), 4) != -1.1282
True
```


Tests for q1_3was1_6 (points: None)
```
>>> pter_over_time.take(0).column(0).item(0) == '1994-01-01'
True
```


Tests for q1_4was1_7 (points: None)
```
>>> highPTER == True or highPTER == False
True
```


Tests for q2_1was2_5 (points: None)
```
>>> type(visualization) == int
True
```
```
>>> 1 <= visualization <= 3
True
```


Tests for q2_3 (points: None)
```
>>> type(assoc) is bool
True
```


Tests for q3_1 (points: None)
```
>>> type(histogram_column_x) == int
True
```
```
>>> histogram_column_x == 1 or histogram_column_x == 2
True
```


Tests for q3_3 (points: None)
```
>>> type(histogram_column_y) == int
True
```
```
>>> histogram_column_y == 1 or histogram_column_y == 2
True
```


Tests for q4_3 (points: None)
```
>>> boston_under_10 >= 0 and boston_under_10 <= 100
True
```
```
>>> manila_under_10 >= 0 and manila_under_10 <= 100
True
```


Tests for q4_4 (points: None)
```
>>> 1 <= manila_median_bin <= 4
True
```


Tests for q5_1was1_1 (points: None)
```
>>> set(['Name', 'Menu_Item', 'Yelp', 'Google', 'Overall', 'Cost']) == set(burritos.labels)
True
```
```
>>> burritos.num_rows == 212
True
```


Tests for q_forbidden (points: None)
```
>>> did_forbidden in [True, False]
True
```
