```python
# Initialize Otter
import otter
grader = otter.Notebook("hw14.ipynb")
```

# Homework 14: Data Ethics: Transparency and Sampling Bias

<!-- BEGIN QUESTION -->

**Question:** What is your name?

_Type your answer here, replacing this text._

<!-- END QUESTION -->



```python
# Don't change this cell; just run it. 

import numpy as np
from datascience import *

# These lines do some fancy plotting magic.
import matplotlib
%matplotlib inline
import matplotlib.pyplot as plt
plt.style.use('fivethirtyeight')
import warnings
warnings.simplefilter('ignore', FutureWarning)
```

## 1. Transparency of Methods and Analysis

### 1.1 Data Dredging

You meet this person who claims to have [telekinetic](https://en.wikipedia.org/wiki/Psychokinesis) powers that enables them to roll a 6 whenever they want on a standard 6-sided die. 
They pull out a die and proceed to roll five 6's in a row.

**Question 1.1.1**
Write a python expression for the probability of rolling five 6's in a row.

```python
five_6s = ...
five_6s
```

```python
grader.check("q1_1_1")
```

**Question 1.1.2** In scientific studies, it is common to say something is [statistically significant](https://en.wikipedia.org/wiki/Statistical_significance) if it has less than 5% chance of having happened randomly.
By this standard, is the fact that this person rolled five 6's statistically significant?
Assign `five6s_significant` to `True` or `False`.

```python
five6s_significant = ...
```

```python
grader.check("q1_1_2")
```

Your friend is intrigued by this result and sets out to find how many students at Ohio University have telekinetic powers.
They get an undergraduate research grant and collect data by having 20,000 students roll five 6-sided dice while concentrating on rolling 6's.

Below we simulate the results.

```python
def one_roll_is_6():
    """simulate the roll of 1 die, returning True if it is a 6 and False otherwise"""
    return np.random.choice(6)==5 # goes 0 to 5, so 5 means 6
# simple test
one_roll_is_6()
```

```python
def n_rolls_are_6s(n):
    """Simulate the rolling of n 6-sided dice, returning True if all are 6's and False otherwise"""
    # roll n times, getting True/False
    nrolls = [one_roll_is_6() for i in np.arange(n)]
    # take the product, get 1 only if all were True
    return np.product(nrolls)
# simple test
n_rolls_are_6s(2)
```

```python
number_of_trials = 20000
rollers_5 = np.array([n_rolls_are_6s(5) for i in np.arange(number_of_trials)])
print("Number who rolled five 6's:",np.count_nonzero(rollers_5))
```

<!-- BEGIN QUESTION -->

**Question 1.1.3** How many students did you find with telekinetic powers?

_Type your answer here, replacing this text._

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

**Question 1.1.4** Your friend is about to post this result on social media and expects to become famous.
Would you advise them to post it or reconsider?
What flaws do you see in their methodology?

_Type your answer here, replacing this text._

<!-- END QUESTION -->

### 1.2 Dissertations

<!-- BEGIN QUESTION -->

**Question 1.2.1.**
Suppose the data you collect contradicts the central claim in your doctoral dissertation.
If you are not an ethical person, what could you do so you still get to graduate?

_Type your answer here, replacing this text._

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

**Question 1.2.2.**
Suppose the data you collect contradicts the central claim in your doctoral dissertation.
If you are an ethical person, what could you do so you still get to graduate?

_Type your answer here, replacing this text._

<!-- END QUESTION -->

### 1.3. Borrowed Scerario

Now we'll consider an [exercise](https://mdsr-book.github.io/mdsr2e/ch-ethics.html#ethics-exercises) from [Modern Data Science with R](https://mdsr-book.github.io/mdsr2e/).
To avoid copyright violations, we will just give the problem number and start.

<!-- BEGIN QUESTION -->

**Question 1.3.1.**

[Problem](https://mdsr-book.github.io/mdsr2e/ch-ethics.html#ethics-exercises) 5 (Medium): A reporter carried out a clinical trial of chocolate ...

Note: The phrase
> different in the treatment group than the outcome. 

should be
> different in the treatment group than the control group. 

_Type your answer here, replacing this text._

<!-- END QUESTION -->

## 2. Sampling Bias

### 2.1 Types of Sampling Bias

<!-- BEGIN QUESTION -->

**Question 2.1.1.**
Suppose that student evaluations of a course only include students who finished the course.
Which type(s) of sampling bias would this introduce?

_Type your answer here, replacing this text._

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

**Question 2.1.2.**
Suppose someone conducts a poll on a local political issue by randomly calling phone numbers in their area code. 
Which type(s) of sampling bias would this introduce?

_Type your answer here, replacing this text._

<!-- END QUESTION -->

### 2.2 Randomized Response

In this section we will simulate the use of a randomized response method in collecting survey data.
The purpose of the method is to:
* increase the accuracy of statistics about things that people may lie about due to embarrassment and/or
* protect the privacy of individual in case the data is leaked or stolen.

Suppose we want to know the answer to a Yes/No question, where people might be embarrassed to answer Yes, but would not be embarrassed to answer No.
For our simulation, we will suppose that for 1/3 of the popuation, the true answer is Yes.
Let's make a table with simulated results.

```python
true_proportion = make_array(1/3,2/3)
survey_size = 10000
survey = Table().with_columns("ID",np.arange(survey_size),
                             "Reality",np.random.choice(['Yes','No'],size=survey_size,p=true_proportion))
survey
```

We expect the proportion of Yes to be about 1/3.

```python
survey.where('Reality','Yes').num_rows/survey_size
```

For randomized response, the person responding needs something random, so we will have each of them flip a fair coin.

```python
survey_with_coin = survey.with_columns("Coin",np.random.choice(['Heads','Tails'],size=survey_size))
survey_with_coin
```

For this basic version of randomized response, the instructions to the survey participant are:
* If you got Heads, answer truthfully.
* If you got Tails, answer Yes.

<!-- BEGIN QUESTION -->

**Question 2.2.1** Why would adding this randomness make people more likely to answer truthfully when they get Heads? 

_Type your answer here, replacing this text._

<!-- END QUESTION -->

**Question 2.2.2** Given that the probability that the reality is Yes is 1/3, what is the probability that the response is Yes? 
Set `response_Yes_probability` to this value.

```python
response_Yes_probability = ...
response_Yes_probability
```

```python
grader.check("q2_2_2")
```

Here is a function that applies the instructions given to the survey participant.

```python
def randomized_response(reality,coin):
    if coin == 'Heads':
        return reality
    return 'Yes'
randomized_response('No','Heads')
```

**Question 2.2.3** Add a column to `survey_with_coin` containing the responses to the survey

```python
response = ...
survey_with_response = survey_with_coin.with_column('Response',response)
survey_with_response
```

```python
grader.check("q2_2_3")
```

Now we can check the number of Yes responses received.

```python
num_yes_responses = survey_with_response.where('Response','Yes').num_rows
num_yes_responses
```

**Question 2.2.4** Based on `num_yes_responses`, what is the simulation's estimate for the proportion of Yes in the population? 

Set `Yes_proportion_estimate` to a python expression that evaluates to this value, using `num_yes_responses` and other quantitites that we have defined.

```python
Yes_proportion_estimate = ...
Yes_proportion_estimate 
```

```python
grader.check("q2_2_4")
```

## Acknowledgements

<!-- BEGIN QUESTION -->

**Acknowledgements:**
If you received help from any person other than your instructor, then describe here what help you received and from whom. 
If not, then say 'None'.

_Type your answer here, replacing this text._

<!-- END QUESTION -->

It is forbidden to: 
* Use the work or programs from students who took this class (in any of its versions at any university).
* Use websites that claim to have solutions for this class (in any of its versions at any university).
* Directly copy from a classmate.

**Did you do anything forbidden while working on this assignment?**

Assign `did_forbidden` to either `True` or `False`.

```python
did_forbidden = ...
```

```python
grader.check("q_forbidden")
```

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit. **Please save before exporting!**

After saving and exporting, submit the .zip file in Canvas

```python
# Save your notebook first, then run this cell to export your submission.
grader.export(run_tests=True)
```

Tests for q1_1_1 (points: None)
```
>>> 0 < five_6s < 1
True
```
```
>>> five_6s < (1 / 6) ** 4
True
```


Tests for q1_1_2 (points: None)
```
>>> five6s_significant in [True, False]
True
```


Tests for q2_2_2 (points: None)
```
>>> 1 / 3 < response_Yes_probability < 1
True
```


Tests for q2_2_3 (points: None)
```
>>> survey_with_response.labels == ('ID', 'Reality', 'Coin', 'Response')
True
```


Tests for q2_2_4 (points: None)
```
>>> 0.2 < Yes_proportion_estimate < 0.5
True
```


Tests for q_forbidden (points: None)
```
>>> did_forbidden in [True, False]
True
```
